# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    import pandas as pd
    mothers_education = pd.read_csv(filepath_or_buffer = 'assets/NISPUF17.csv' ,  index_col = None, header = 0, usecols = ['EDUC1', 'SEQNUMC'])

    #Find total number of records 
    mothers_education['EDUC1'].replace(to_replace = {1: "less than high school", 2:"high school", 3:"more than high school but not college", 4: "college"}, inplace = True)

    percentages = mothers_education.groupby(by = 'EDUC1', axis = 0).count()/mothers_education['SEQNUMC'].count()
    #print(percentages)
    #print(type(percentages.T))
    returnable = percentages['SEQNUMC'].to_dict()
    #returnable    

    return (returnable)
    raise NotImplementedError()
```

```python
    import pandas as pd
    mothers_education = pd.read_csv(filepath_or_buffer = 'assets/NISPUF17.csv' ,  index_col = None, header = 0, usecols = ['EDUC1', 'SEQNUMC'])

    #Find total number of records 
    mothers_education['EDUC1'].replace(to_replace = {1: "less than high school", 2:"high school", 3:"more than high school but not college", 4: "college"}, inplace = True)
    
    percentages = mothers_education.groupby(by = 'EDUC1', axis = 0).count()/mothers_education['SEQNUMC'].count()
    print(percentages)
    returnable = percentages['SEQNUMC'].to_dict()
    returnable
    #dict(percentages)
    #print(type(percentages))
    #print(percentages)
    #print(type(percentages.T))
    #percentages = percentages.to_dict(orient="index")
    #print(percentages)
```

Output:
```
SEQNUMC
EDUC1                                          
college                                0.479747
high school                            0.172352
less than high school                  0.102020
more than high school but not college  0.245881
{'college': 0.47974705779026877,
 'high school': 0.172352011241876,
 'less than high school': 0.10202002459160373,
 'more than high school but not college': 0.24588090637625154}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    import pandas as pd
    df= pd.read_csv(filepath_or_buffer = 'assets/NISPUF17.csv' ,  index_col = None, header = 0, usecols =['CBF_01', 'P_NUMFLU'] )
    #print(df.head(5))
    #Filter the data to only those children who have been known to have been breast fed or known not have been breast fed
    mask = (df['CBF_01']==1) | ( df['CBF_01']==2)
    mask.head(5)
    df_selected = df[mask]
    #df.reset_index()
    #print(df)
    #Find mean number of influenze vaccines for both groups
    series = df_selected.groupby(by = 'CBF_01').mean() 
    
     list1=   [x  for x in series['P_NUMFLU']]
    tuple1 = tuple(list1)
    print(tuple1)   
    #print(series.to_dict(list))
    return(tuple1)
    raise NotImplementedError()
```

```python
average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
tuple
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

Output:
```
CBF_01           1         2
P_NUMFLU  1.879919  1.596395
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    #HAD_CPOX --did the child ever have chicken pox, Yes (1), No(2), Others
    #P_NUMVRC - Total number of varicella doses, 0, 1, 2, NA
    # SEx - Sex of the child , Male (1), Female (2)

    df= pd.read_csv(filepath_or_buffer = 'assets/NISPUF17.csv' ,  index_col = None, header = 0, usecols =['SEX', 'HAD_CPOX', 'P_NUMVRC'] )
    df_male = df [df['SEX']==1]

    df_male_cleaned =df_male.dropna()
    mask =  (df_male_cleaned['P_NUMVRC']> 0.0) & (df_male_cleaned['HAD_CPOX'].isin([1,2]))
    df_male_vaccinated = df_male_cleaned[mask]

    import numpy as np

    df_male_vaccinated['infected'] =  list(np.where( df_male_vaccinated['HAD_CPOX'] ==1  ,'Infected at least once','Uninfected') )
    #print(df_male_vaccinated)
    df_male_vaccinated.drop(['SEX','P_NUMVRC'], axis = 1, inplace =True)
    df_male_grouped = df_male_vaccinated.groupby(by = ['infected']).count()
    df_male_grouped
    dict = {}
    dict['male'] = df_male_grouped.loc['Infected at least once']['HAD_CPOX']/df_male_grouped.loc['Uninfected']['HAD_CPOX']
    #dict 

    df_female = df [df['SEX']==2]

    #print(df_male.head(5))
    df_female_cleaned =df_female.dropna()
    mask =  (df_female_cleaned['P_NUMVRC']> 0.0) & (df_female_cleaned['HAD_CPOX'].isin([1,2]))
    df_female_vaccinated = df_female_cleaned[mask]

   # import numpy as np
    df_female_vaccinated['infected'] =  list(np.where( df_female_vaccinated['HAD_CPOX'] ==1  ,'Infected at least once','Uninfected') )
    #print(df_female_vaccinated)
    df_female_vaccinated.drop(['SEX','P_NUMVRC'], axis = 1, inplace =True)
    df_female_grouped = df_female_vaccinated.groupby(by = ['infected']).count()
    df_female_grouped
    dict['female'] = df_female_grouped.loc['Infected at least once']['HAD_CPOX']/df_female_grouped.loc['Uninfected']['HAD_CPOX']
    print(dict)
    return(dict)
    raise NotImplementedError()
```

```python
len(chickenpox_by_sex())
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
C:\Users\singl\AppData\Local\Temp/ipykernel_16436/3245271903.py:16: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  df_male_vaccinated['infected'] =  list(np.where( df_male_vaccinated['HAD_CPOX'] ==1  ,'Infected at least once','Uninfected') )
C:\Users\singl\AppData\Local\Temp/ipykernel_16436/3245271903.py:33: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  df_female_vaccinated['infected'] =  list(np.where( df_female_vaccinated['HAD_CPOX'] ==1  ,'In
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
C:\Users\singl\AppData\Local\Temp/ipykernel_16436/3245271903.py:16: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  df_male_vaccinated['infected'] =  list(np.where( df_male_vaccinated['HAD_CPOX'] ==1  ,'Infected at least once','Uninfected') )
C:\Users\singl\AppData\Local\Temp/ipykernel_16436/3245271903.py:33: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  df_female_vaccinated['infected'] =  list(np.where( df_female_vaccinated['HAD_CPOX'] ==1  ,'In
```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    #df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
     #              "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})

    # here is some stub code to actually run the correlation
    #corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    df  = pd.read_csv(filepath_or_buffer = 'assets/NISPUF17.csv' ,  index_col = None, header = 0, usecols = ['HAD_CPOX', 'P_NUMVRC'])
    df.head(5)
    selected_filtered_df = df[df['HAD_CPOX'].isin([1,2])].dropna()

    corr, pval=stats.pearsonr(selected_filtered_df["HAD_CPOX"],selected_filtered_df["P_NUMVRC"])
    #print(corr)
    #print (pval)
    
    # just return the correlation
    return corr

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python

```