# Assignment 1: Regular Expressions (Regex)

## Overview
This assignment focuses on **Regular Expressions (Regex)**, a powerful text-pattern matching tool built into Python's `re` module. Regex is a fundamental data science skill for parsing and cleaning unstructured text data.

## Learning Objectives
- Understand regex pattern syntax: character classes, anchors, quantifiers, lookaheads
- Use `re.findall()` and `re.finditer()` to extract matches from text
- Apply regex to realistic tasks: name extraction, grade filtering, and log file parsing

## Assignment Structure
| Part | Task | Key Regex Concepts |
|------|------|--------------------|
| Part A | Extract names from a string | Character classes `[A-Z][a-z]*` |
| Part B | Filter students with grade "B" from a file | Lookahead assertions `(?=...)` |
| Part C | Parse a web server access log into dictionaries | Named groups, multiple patterns |

## Data Files
- `assets/grades.txt` — Line-separated list of students and their grades
- `assets/logdata.txt` — Standard Apache/Nginx web server access log (979 entries)

## Instructions
1. Write your solution **inside** each function body
2. Each function must **return** the result
3. Do not modify the `assert` statements — they validate your output automatically
4. You may use regex documentation and cheat sheets as reference

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

## Example Walkthrough

The example below shows the expected pattern: define a function, write logic inside, and return the result. The autograder calls your function and checks the return value.

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    # Match capitalized words — each starts with an uppercase letter followed by lowercase letters
    pattern = "[A-Z][a-z]*"
    result = re.findall(pattern, simple_string)
    return result
```

## Part A

Find a list of all of all of the names in the following string using regex.

```python
import re
def grades():
    with open("assets/grades.txt", "r") as file:
        grades = file.read()

    # Match "First Last" names where the grade immediately after is B
    # The lookahead (?=: B) ensures we only get students who received a B
    pattern = r"[A-Z][a-z]* [A-Z][a-z]*(?=: B)"
    result = re.findall(pattern, grades)
    return result
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # Four separate patterns for each field in the Apache Common Log Format
    host_pattern = r"\d{1,3}\.\d{1,3}\.\d{1,3}\.\d{1,3}"
    user_pattern = r"(?<=- )[\w\-]+(?= \[)"
    time_pattern = r"(?<=\[)[^\]]+(?=\])"
    request_pattern = r'(?<=")[^"]+(?=")'
    
    hosts = re.findall(host_pattern, logdata)
    users = re.findall(user_pattern, logdata)
    times = re.findall(time_pattern, logdata)
    requests = re.findall(request_pattern, logdata)
    
    # Combine the four lists into a list of dictionaries
    result = []
    for i in range(len(users)):
        result.append({
            'host': hosts[i],
            'user_name': users[i],
            'time': times[i],
            'request': requests[i]
        })
    return result
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # YOUR CODE HERE
    raise NotImplementedError()
```