# Assignment 2: CDC Immunization Data Analysis

## Overview
This assignment analyzes **2017 CDC National Immunization Survey (NIS-PUF17)** data to explore relationships between maternal education, breastfeeding, and childhood vaccination patterns. It builds foundational pandas data manipulation skills.

## Learning Objectives
- Load and filter CSV data with `pandas.read_csv()`
- Use `groupby()` and aggregation functions to compute group statistics
- Calculate proportions, ratios, and Pearson correlations
- Work with real public health survey data

## Dataset
- **File**: `assets/NISPUF17.csv` — 2017 CDC National Immunization Survey, Public-Use File
- **Reference**: `assets/NIS-PUF17-DUG.pdf` — Data Users Guide (variable definitions)

### Key Variables
| Variable | Description |
|----------|-------------|
| `EDUC1` | Mother's education level (1=<HS, 2=HS, 3=>HS, 4=College) |
| `CBF_01` | Child breastfed (1=Yes, 2=No) |
| `P_NUMFLU` | Number of influenza vaccine doses received |
| `HAD_CPOX` | Child had chickenpox (1=Yes, 2=No) |
| `P_NUMVRC` | Number of varicella (chickenpox) vaccine doses |
| `SEX` | Child's sex (1=Male, 2=Female) |

## Instructions
- Write solutions inside each function body and `return` the result
- Do not modify the assert statements
- Use `pandas` for all data manipulation

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

### Approach
Load the CSV keeping only `EDUC1` and `SEQNUMC` columns. Remap the numeric education codes to descriptive strings, then use `groupby().count()` divided by total count to get proportions. Convert to a dictionary.

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    import pandas as pd
    df = pd.read_csv('assets/NISPUF17.csv', index_col=None, header=0, usecols=['EDUC1', 'SEQNUMC'])
    
    # Map numeric education codes to descriptive labels
    df['EDUC1'].replace({
        1: "less than high school",
        2: "high school",
        3: "more than high school but not college",
        4: "college"
    }, inplace=True)
    
    # Calculate proportion of children in each education group
    total = df['SEQNUMC'].count()
    proportions = df.groupby('EDUC1')['SEQNUMC'].count() / total
    return proportions.to_dict()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

Output:
```
C:\Users\singl\AppData\Local\Temp\ipykernel_21676\1911311424.py:6: FutureWarning: A value is trying to be set on a copy of a DataFrame or Series through chained assignment using an inplace method.
The behavior will change in pandas 3.0. This inplace method will never work because the intermediate object on which we are setting values always behaves as a copy.

For example, when doing 'df[col].method(value, inplace=True)', try using 'df.method({col: value}, inplace=True)' or df[col] = df[col].method(value) instead, to perform the operation inplace on the original object.


  df['EDUC1'].replace({
C:\Users\singl\AppData\Local\Temp\ipykernel_21676\1911311424.py:6: FutureWarning: A value is trying to be set on a copy of a DataFrame or Series through chained assignment using an inplace method.
The behavior will change in pandas 3.0. This inplace method will never work because the intermediate object on which we are setting values always behaves as a copy.

For example, when doing 'df[col].m
```

### Approach
Load `CBF_01` (breastfed indicator) and `P_NUMFLU` (flu vaccine count). Filter to only known values (1=breastfed, 2=not breastfed), then use `groupby().mean()` to get the average flu doses for each group. Return as a tuple (breastfed_avg, not_breastfed_avg).

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    import pandas as pd
    df = pd.read_csv('assets/NISPUF17.csv', index_col=None, header=0, usecols=['CBF_01', 'P_NUMFLU'])
    
    # Keep only children with known breastfeeding status (1=Yes, 2=No)
    df_known = df[df['CBF_01'].isin([1, 2])]
    
    # Calculate average flu vaccine doses for each group
    averages = df_known.groupby('CBF_01')['P_NUMFLU'].mean()
    return tuple(averages)
```

### Approach
This question asks for the ratio of vaccinated-but-infected children vs vaccinated-but-not-infected, broken down by sex. Filter to children who: (a) received at least one varicella vaccine dose, AND (b) have a known chickenpox status (1 or 2). Then group by sex and compute: count(had_cpox=1) / count(had_cpox=2) for each sex.

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

### Approach
Load `HAD_CPOX` and `P_NUMVRC`. Filter to rows where `HAD_CPOX` is 1 or 2 (excluding unknowns) and drop NaN values. Use `scipy.stats.pearsonr()` to compute the Pearson correlation coefficient between the two columns. A negative correlation indicates that children who had chickenpox tend to have fewer vaccine doses, which is the expected biological relationship.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    import pandas as pd
    import numpy as np
    df = pd.read_csv('assets/NISPUF17.csv', index_col=None, header=0,
                     usecols=['SEX', 'HAD_CPOX', 'P_NUMVRC'])
    
    # Keep vaccinated children (>=1 dose) with known chickenpox status
    df = df.dropna()
    df = df[(df['P_NUMVRC'] >= 1) & (df['HAD_CPOX'].isin([1, 2]))]
    
    result = {}
    for sex_code, sex_label in [(1, 'male'), (2, 'female')]:
        group = df[df['SEX'] == sex_code]
        infected = (group['HAD_CPOX'] == 1).sum()
        not_infected = (group['HAD_CPOX'] == 2).sum()
        result[sex_label] = infected / not_infected
    
    return result
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    df = pd.read_csv('assets/NISPUF17.csv', index_col=None, header=0,
                     usecols=['HAD_CPOX', 'P_NUMVRC'])
    
    # Keep only rows where chickenpox status is known (1=Yes, 2=No) and drop missing vaccine data
    df = df[df['HAD_CPOX'].isin([1, 2])].dropna()
    
    # Compute Pearson correlation between chickenpox status and number of vaccine doses
    corr, pval = stats.pearsonr(df['HAD_CPOX'], df['P_NUMVRC'])
    return corr
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```