# Course 5 - Applied Social Network Analysis in Python
## Assignment 2: Network Connectivity in a Company Email Graph

### Overview
This assignment analyzes a directed multigraph of internal company emails to study network connectivity, connected components, graph diameter, and network flow. Understanding connectivity helps determine how information propagates through an organization.

### Learning Objectives
- Load and query directed multigraphs with edge timestamps
- Distinguish between strong and weak connectivity
- Find largest connected components and extract subgraphs
- Compute graph distance metrics: average path length, diameter, eccentricity
- Apply node cut algorithms to find minimum cuts in directed graphs

### Dataset
- **File:** `assets/email_network.txt` — Directed multigraph edge list
- **Format:** `sender recipient timestamp` (tab-delimited)
- **Size:** 167 nodes (employees), 82,927 emails (edges)
- **Edge attribute:** `time` (Unix timestamp float)

### Assignment Structure
| Question | Task |
|----------|------|
| Q1 | Load as directed MultiDiGraph with time edge attribute |
| Q2 | Count employees (nodes) and unique sender-recipient pairs |
| Q3 | Is information flow possible between all pairs? (strongly vs weakly connected) |
| Q4 | Size of largest weakly connected component |
| Q5 | Size of largest strongly connected component |
| Q6 | Extract subgraph G_sc of the largest strongly connected component |
| Q7 | Average shortest path length in G_sc |
| Q8 | Diameter (maximum shortest path) of G_sc |
| Q9 | Nodes with eccentricity equal to the diameter (peripheral nodes) |
| Q10 | Nodes with eccentricity equal to the radius (central nodes) |
| Q11 | Node with most shortest paths of length = diameter; how many such paths |
| Q12 | Minimum node cut to block flow from node 97 to node 10 |
| Q13 | Convert G_sc to undirected graph G_un |
| Q14 | Transitivity and average clustering coefficient of G_un |

### Key Concepts
- **Strong connectivity:** Every node can reach every other node (directed paths)
- **Weak connectivity:** Underlying undirected graph is connected
- **Eccentricity:** Maximum shortest path distance from a node to any other node
- **Diameter:** Maximum eccentricity; **Radius:** Minimum eccentricity

```python
import networkx as nx

#!head assets/email_network.txt
```

### What This Code Does
Imports NetworkX for graph analysis. The email network is stored in a text file with three columns: sender, recipient, and Unix timestamp. It will be loaded as a `MultiDiGraph` (directed graph allowing multiple edges between the same pair of nodes) because the same person may send multiple emails to the same recipient at different times.

### Question 1

Using networkx, load up the directed multigraph from `assets/email_network.txt`. Make sure the node names are strings.

*This function should return a directed multigraph networkx graph.*

<!-- grade_id: cell-60316968695c6f20, points: 0, grade: false, solution: true -->
```python
import networkx as nx
def answer_one():
    
    # YOUR CODE HERE
        # YOUR CODE HERE
    path = "assets/email_network.txt"
    #G= nx.read_edgelist()

    G = nx.read_edgelist(path=path, data=[ ('time', float)], 
                            create_using=nx.MultiDiGraph())

    G.is_directed(), G.is_multigraph()
    #display (G.nodes())
    #display(list(G.edges(data = True)))
    #display(G.edges())
    return G
    raise NotImplementedError()
```

<!-- grade_id: cell-483a56c89f9db231, points: 1, grade: true, solution: false -->
```python
ans_one = answer_one()
```

### Question 2

How many employees are represented in the network?

How many `sender`->`recipient` pairs of employees are there in the network such that `sender` sent at least one email to `recipient`? Note that even if a `sender` sent multiple messages to a `recipient`, they should only be counted once. You should **not** exclude cases where an employee sent emails to themselves from this [email] count. 

*This function should return a tuple with two integers (#employees, # `sender`->`recipient` pairs).*

<!-- grade_id: cell-16a78f22207d9f95, points: 0, grade: false, solution: true -->
```python
def answer_two():
    # YOUR CODE HERE
    G = answer_one()
    employee_count = len(G.nodes())

    display(print(employee_count))

    connected_nodes= len(set(G.edges()))
    display(print(connected_nodes))
    #raise NotImplementedError()
    return (employee_count, connected_nodes)
    raise NotImplementedError()
```

<!-- grade_id: cell-5b6391549b076d2b, points: 1, grade: true, solution: false -->
```python
ans_two = answer_two()

```

Output:
```
167
None
5784
None
```

### Question 3

* Part 1. Assume that information in this company can only be exchanged through email.

    When an employee sends an email to another employee, a communication channel has been created, allowing the sender to provide information to the reciever, but not viceversa. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


* Part 2. Now assume that a communication channel established by an email allows information to be exchanged both ways. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


*This function should return a tuple of bools (part1, part2).*

<!-- grade_id: cell-606dd37fb2f55358, points: 0, grade: false, solution: true -->
```python
def answer_three():
    # YOUR CODE HERE
    G=answer_one()
    a  =nx.is_strongly_connected(G)
    b =nx.is_weakly_connected(G)
    #raise NotImplementedError()
    return a,b    
    raise NotImplementedError()
```

<!-- grade_id: cell-82b3f0bc45e2895f, points: 1, grade: true, solution: false -->
```python
ans_three = answer_three()

```

### Question 4

How many nodes are in the largest weakly connected component of the graph?

*This function should return an int.*

<!-- grade_id: cell-39196f4c27b35534, points: 0, grade: false, solution: true -->
```python
def answer_four():
    # YOUR CODE HERE
    G = answer_one()
    a=sorted(nx.weakly_connected_components(G))
    b=len(a[0])
    print(b)
    #raise NotImplementedError()
    return b    
    raise NotImplementedError()
```

<!-- grade_id: cell-2b1b7b06ecfa751d, points: 1, grade: true, solution: false -->
```python
ans_four = answer_four()
```

Output:
```
167
```

### Question 5

How many nodes are in the largest strongly connected component?

*This function should return an int*

<!-- grade_id: cell-879d17a9db5c3790, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    G= answer_one()
    component_list = sorted(nx.strongly_connected_components(G))
    component_list

    sorted_list = sorted(component_list, key = len, reverse = True)
    len_longest = len(sorted_list[0])
    len_longest
    #raise NotImplementedError()
    return len_longest    
    raise NotImplementedError()
```

<!-- grade_id: cell-b0524f7dc1fbdec4, points: 1, grade: true, solution: false -->
```python
ans_five = answer_five()

```

### Question 6

Using the NetworkX functions `strongly_connected_components` and `subgraph`, find the subgraph of nodes in the largest strongly connected component. 
Call this graph G_sc.

*This function should return a networkx MultiDiGraph named G_sc.*

<!-- grade_id: cell-a81e6f88ef84578c, points: 0, grade: false, solution: true -->
```python
def answer_six():
    # YOUR CODE HERE
    G = answer_one()
    component_list = sorted(nx.strongly_connected_components(G))
    component_list

    sorted_list = sorted(component_list, key = len, reverse = True)
    longest = sorted_list[0]
    longest

    G_sc=nx.subgraph(G, nbunch = longest)
    #raise NotImplementedError()
    return G_sc
    raise NotImplementedError()
```

<!-- grade_id: cell-cf148ef273b3b19c, points: 1, grade: true, solution: false -->
```python
ans_six = answer_six()
assert type(ans_six) == nx.MultiDiGraph , "Your return type should be a MultiDiGraph object"

```

### Question 7

What is the average distance between nodes in G_sc?

*This function should return a float.*

<!-- grade_id: cell-1a3075039b3c2c0b, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    G_sc = answer_six()

    a= nx.average_shortest_path_length(G_sc)
    return a    
    raise NotImplementedError()
```

<!-- grade_id: cell-5b374fdd48f37e02, points: 1, grade: true, solution: false -->
```python
ans_seven = answer_seven()

```

### Question 8

What is the largest possible distance between two employees in G_sc?

*This function should return an int.*

<!-- grade_id: cell-4fbb8d34d00f675c, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    # YOUR CODE HERE
    G_sc = answer_six()

    a=nx.diameter(G_sc)
    return a    
    raise NotImplementedError()
```

<!-- grade_id: cell-c5714787854ef644, points: 1, grade: true, solution: false -->
```python
ans_eight = answer_eight()

```

### Question 9

What is the set of nodes in G_sc with eccentricity equal to the diameter?

*This function should return a set of the node(s).*

<!-- grade_id: cell-bcfab850744277ed, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    G_sc = answer_six()
    nodes_eccentricity_d = nx.eccentricity(G_sc)
    type(nodes_eccentricity_d)

    a=[key for key in nodes_eccentricity_d.keys() if nodes_eccentricity_d[key] ==nx.diameter(G_sc)]
    b = set(a)
    b
    return b    
    raise NotImplementedError()
```

<!-- grade_id: cell-77c9ca0b94df3d6f, points: 1, grade: true, solution: false -->
```python
ans_nine = answer_nine()
assert type(ans_nine) == set, "Student answer must return a set"
```

### Question 10

What is the set of node(s) in G_sc with eccentricity equal to the radius?

*This function should return a set of the node(s).*

<!-- grade_id: cell-fd3d43ea37da03f7, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    G_sc = answer_six()

    nodes_eccentricity_d = nx.eccentricity(G_sc)
    type(nodes_eccentricity_d)
    a=[key for key in nodes_eccentricity_d.keys() if nodes_eccentricity_d[key] ==nx.radius(G_sc)]
    b = set(a)
    b
    return b    
    raise NotImplementedError()
```

<!-- grade_id: cell-bfd2ee304bc25264, points: 1, grade: true, solution: false -->
```python
ans_ten = answer_ten()
assert type(ans_ten) == set, "Student answer must return a set"
```

### Question 11

Which node in G_sc has the most shortest paths to other nodes whose distance equal the diameter of G_sc?


For the node with the most such shortest paths, how many of these paths are there?


*This function should return a tuple (name of node, number of paths).*

<!-- grade_id: cell-c40858baf3d92139, points: 0, grade: false, solution: true -->
```python
# for all nodes A:
#     A_path_count = 0
#     for all nodes B not equal to A:
#        if the shortest path from A to B is equal the diameter:
#            A_path_count += number of shortest paths from A to B
#     keep (A, A_path_count) if is has the largest A_path_count
# return (A, A_path_count) for the node A with the largest A_path_count

def answer_eleven():
    # YOUR CODE HERE
    #from icecream import ic

    G_sc = answer_six()

    d_Gsc = nx.diameter(G_sc)
    display(d_Gsc)

    node_list = list(G_sc.nodes())
                    
    #display(node_list)

    node_list_out = []
    for node in node_list:
        path_cnt = 0 
        revised_node_list = [ target_node for target_node in node_list if target_node not in  [node]]

        #display(revised_node_list)
        for target_node in revised_node_list:
            if target_node ==node:
                pass
            else:
                shortest_path_length = nx.shortest_path_length(G_sc, source =node, target = target_node ) 
                #ic(shortest_path_length)
                
                if shortest_path_length == d_Gsc:
                    path_cnt = path_cnt+1
                    #display("pth_cnt_incremented")
        node_list_out.append((node, path_cnt))

        #ic(node_list_out)

    display(node_list_out)

    node_list_sorted = sorted(node_list_out, key = lambda x: x[1], reverse = True)
    a=node_list_sorted[0]
    return a    
    raise NotImplementedError()
```

<!-- grade_id: cell-f79b06650f61cf37, points: 1, grade: true, solution: false -->
```python
ans_eleven = answer_eleven()
assert type(ans_eleven) == tuple, "Student answer must be a tuple"
```

Output:
```
3
[('1', 0),
 ('2', 0),
 ('3', 0),
 ('4', 0),
 ('5', 0),
 ('6', 0),
 ('7', 0),
 ('8', 0),
 ('9', 0),
 ('10', 0),
 ('11', 0),
 ('12', 0),
 ('13', 0),
 ('14', 0),
 ('15', 0),
 ('16', 0),
 ('17', 0),
 ('18', 0),
 ('19', 0),
 ('20', 0),
 ('21', 0),
 ('22', 0),
 ('23', 0),
 ('24', 0),
 ('25', 0),
 ('26', 0),
 ('27', 0),
 ('28', 0),
 ('29', 0),
 ('30', 0),
 ('31', 0),
 ('32', 0),
 ('33', 0),
 ('34', 0),
 ('35', 0),
 ('36', 0),
 ('37', 0),
 ('38', 0),
 ('39', 0),
 ('40', 0),
 ('41', 0),
 ('42', 0),
 ('43', 0),
 ('44', 0),
 ('45', 0),
 ('46', 0),
 ('47', 0),
 ('48', 0),
 ('49', 0),
 ('50', 0),
 ('51', 0),
 ('52', 0),
 ('53', 0),
 ('54', 0),
 ('55', 0),
 ('56', 0),
 ('57', 0),
 ('58', 0),
 ('59', 0),
 ('60', 0),
 ('61', 0),
 ('62', 0),
 ('63', 0),
 ('64', 0),
 ('65', 0),
 ('66', 0),
 ('67', 0),
 ('68', 0),
 ('69', 0),
 ('70', 0),
 ('71', 0),
 ('72', 0),
 ('73', 0),
 ('74', 0),
 ('75', 0),
 ('76', 0),
 ('77', 0),
 ('78', 0),
 ('79', 0),
 ('80', 0),
 ('81', 0),
 ('82', 0),
 ('83', 0),
 ('84', 0),
```

### Question 12

Suppose you want to prevent communication flow from the node that you found in question 11 to node 10. What is the smallest number of nodes you would need to remove from the graph (you're not allowed to remove the node from the previous question or 10)? 

*This function should return an integer.*

<!-- grade_id: cell-a4606db8971dcbba, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    G_sc = answer_six()
    a =nx.minimum_node_cut(G_sc, '97', '10')
    return a    
    raise NotImplementedError()
```

<!-- grade_id: cell-509cfa9f4136124d, points: 1, grade: true, solution: false -->
```python
ans_twelve = answer_twelve()

```

### Question 13

Convert the graph G_sc into an undirected graph by removing the direction of the edges of G_sc. Call the new graph G_un. 


*This function should return a networkx Graph.*

<!-- grade_id: cell-612b40425f0851e8, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    G_sc = answer_six()
    G_un = nx.Graph(G_sc)
    display(type(G_un))
    return G_un    
    raise NotImplementedError()
```

<!-- grade_id: cell-d1c0627a327cd774, points: 1, grade: true, solution: false -->
```python
ans_thirteen = answer_thirteen()
assert type(ans_thirteen) == nx.Graph , "Your return type should be a Graph object"

```

Output:
```
networkx.classes.graph.Graph
```

### Question 14

What is the transitivity and average clustering coefficient of graph G_un?

*This function should return a tuple (transitivity, avg clustering).*     
*Note: DO NOT round up your answer.*

<!-- grade_id: cell-63e48b1eed1a1fb8, points: 0, grade: false, solution: true -->
```python
def answer_fourteen():
    # YOUR CODE HERE
    G_un = answer_thirteen()
    a=nx.average_clustering(G_un)

    b=nx.transitivity(G_un)
    return (b,a)    
    raise NotImplementedError()
```

<!-- grade_id: cell-41dda91202f58e7e, points: 1, grade: true, solution: false -->
```python
ans_fourteen = answer_fourteen()
assert type(ans_fourteen) == tuple, "Student answer must be a tuple"

```

Output:
```
networkx.classes.graph.Graph
```

```python

```